# Power and Effect Size Checkpoint

```python
# run this cell without changes

import numpy as np
import pandas as pd
from statsmodels.stats import power
from scipy import stats
```

#### 1. What is the relationship between power and the false-negative rate, $\beta$ ?

<!-- grade_id: 922a8ab92816117f6fb5a8a716c38719, points: 1, grade: true, solution: true -->
=== BEGIN MARK SCHEME ===

Power is simply 1 less the false-negative rate, 1 - beta.

=== END MARK SCHEME ===

#### 2. Calculate Cohen's *d* for the following data.  

Suppose we have the following samples of ant body lengths in centimeters, taken from two colonies that are distinct but spatially close to one another.

Calculate Cohen's *d* for these groups and assign it to the variable `d`.

(The variance of each group has been provided for you.)

```python
# run this cell without changes

colony1 = np.array([2.101, 2.302, 2.403])
colony2 = np.array([2.604, 2.505, 2.506])

colony1_var = np.var(colony1, ddof=1)
colony2_var = np.var(colony2, ddof=1)
```

<!-- grade_id: cbbca1a5c0b5d7f22c660022597d0f7d, points: 0, grade: false, solution: true -->
```python
# Replace None with appropriate code
d = None
### BEGIN SOLUTION

n1, n2 = len(colony1), len(colony2)

pooled_stdev = np.sqrt(
    ((n1-1)*colony1_var + (n2-1)*colony2_var) 
    / 
    (n1 + n2 - 2)
)

d = (colony2.mean() - colony1.mean()) / pooled_stdev

def Cohen_d_1(group1, group2):
    """
    This function comes from the Effect Size lab
    """

    # Compute Cohen's d.

    # group1: Series or NumPy array
    # group2: Series or NumPy array

    # returns a floating point number 

    diff = group1.mean() - group2.mean()

    n1, n2 = len(group1), len(group2)
    var1 = group1.var()
    var2 = group2.var()

    # Calculate the pooled threshold as shown earlier
    pooled_var = (n1 * var1 + n2 * var2) / (n1 + n2)
    
    # Calculate Cohen's d statistic
    d = diff / np.sqrt(pooled_var)
    
    return d

def Cohen_d_2(group1, group2):
    """
    This function comes from Greg's lecture
    """

    """
    Computes Cohen's d.
    """
    
    # group1: Series or NumPy array
    # group2: Series or NumPy array

    # returns a floating point number 

    diff = group1.mean() - group2.mean()

    n1 = len(group1)
    n2 = len(group2)
    var1 = group1.var(ddof=1)
    var2 = group2.var(ddof=1)

    # Calculate the pooled variance
    pooled_var = ((n1-1) * var1 + (n2-1) * var2) / (n1 + n2 - 2)
    
    # Calculate Cohen's d statistic
    d = diff / np.sqrt(pooled_var)
    
    return d
### END SOLUTION

print(d)
```

Output:
```
2.3265865171551505
```

<!-- grade_id: bc76efa02a9d924944d8601cb7298908, points: 1, grade: true, solution: false -->
```python
# PUT ALL WORK FOR THE ABOVE QUESTION ABOVE THIS CELL
# THIS UNALTERABLE CELL CONTAINS HIDDEN TESTS

# d should be a floating point number
assert type(d) == float or type(d) == np.float64

### BEGIN HIDDEN TESTS

# Still give them points for negative d value
d = np.abs(d)

# Setup variables for tests
n1, n2 = len(colony1), len(colony2)
standardizer_1 = (n1 * colony1_var + n2 * colony2_var) / (n1 + n2)
standardizer_2 = ((n1-1) * colony1_var + (n2-1) * colony2_var) / (n1 + n2 - 2)

assert d == (colony2.mean() - colony1.mean()) / np.sqrt(standardizer_1) or \
       d == (colony2.mean() - colony1.mean()) / np.sqrt(standardizer_2)


### END HIDDEN TESTS
```

#### 3. Is this a large effect size? What conclusion should we draw about these two samples?

<!-- grade_id: 07f8b7b210327354d2a14dcb9b52decf, points: 1, grade: true, solution: true -->
=== BEGIN MARK SCHEME ===

This is a huge effect size, but we have to be careful about drawing conclusions about the relationship between the two colonies (that the colonies contain distinct populations of ants, for example), since we have such small sample sizes.

=== END MARK SCHEME ===

#### 4. We decide we want to collect more data to have a more robust experiment. 

#### Given the current effect size, calculate how many observations we need (of each ant colony) if we want our test to have a power rating of 0.9 and a false-positive rate of 5%.

<!-- grade_id: 7cde4277651e3aff7fb9b663b58103b2, points: 0, grade: false, solution: true -->
```python
# Replace None with appropriate code
observations_needed = None
### BEGIN SOLUTION
from test_scripts.test_class import Test
test = Test()

observations_needed = power.TTestIndPower().solve_power(effect_size=d,
                                  alpha=0.05,
                                 power=0.9)

test.save(observations_needed, "observations_needed")

### END SOLUTION
print(observations_needed)
```

Output:
```
5.066906633232539
```

<!-- grade_id: 5afc4e42316bdceb2ee45195f1874697, points: 1, grade: true, solution: false -->
```python
# PUT ALL WORK FOR THE ABOVE QUESTION ABOVE THIS CELL
# THIS UNALTERABLE CELL CONTAINS HIDDEN TESTS

# No need to cast the result to an int. observations_needed should be a floating point number
assert type(observations_needed) == float or type(observations_needed) == np.float64

### BEGIN HIDDEN TESTS

from test_scripts.test_class import Test
test = Test()

test.run_test(observations_needed, "observations_needed")

### END HIDDEN TESTS
```

Suppose we gather more data on our ants and then re-visit our calculations. 

```python
# run this cell without changes

col1 = pd.read_csv('data/colony1', index_col=0)
col2 = pd.read_csv('data/colony2', index_col=0)
```

```python
# run this cell without changes
col1.head()
```

Output:
```
col1_length
0        2.101
1        2.302
2        2.403
3        3.341
4        2.626
```

```python
# run this cell without changes
col2.head()
```

Output:
```
col2_length
0     2.604000
1     2.505000
2     2.506000
3     2.383960
4     2.578711
```

#### 5. Do a two-tailed t-test on the lengths of ants in these two colonies.

<!-- grade_id: 8df29a4abc2d72d26a2f258089e3f7a9, points: 0.5, grade: true, solution: true -->
```python
### BEGIN SOLUTION

stats.ttest_ind(col1['col1_length'], col2['col2_length'])

### END SOLUTION
```

What should we conclude about these two collections of ant body lengths?

<!-- grade_id: f588d67ce53c6f183eb9e8910de453c2, points: 0.5, grade: true, solution: true -->
=== BEGIN MARK SCHEME ===

The p-value is now quite large, and so we cannot reject the null hypothesis that says that there is no difference between the body lengths in the two ant colonies.

=== END MARK SCHEME ===